```python
# Initialize Otter
import otter
grader = otter.Notebook("hw4.ipynb")
```

# CPSC 330 - Applied Machine Learning 

## Homework 4: Logistic regression, hyperparameter optimization 
### Associated lectures: [Lectures 7, 8](https://github.com/UBC-CS/cpsc330-2023s/tree/main/lectures) 

**Due date: See the [Calendar](https://calendar.google.com/calendar/u/0/embed?src=7a04205ae91b85e82ebc74daddbf2933c6b6723b81abb966f0e69c66a996c43b@group.calendar.google.com&ctz=America/Vancouver&pli=1).**

## Imports 

```python
from hashlib import sha1
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd

plt.rcParams["font.size"] = 16

from sklearn.dummy import DummyClassifier
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import (
    GridSearchCV,
    cross_val_score,
    cross_validate,
    train_test_split,
)
from sklearn.pipeline import Pipeline, make_pipeline
from sklearn.tree import DecisionTreeClassifier
```

## Instructions
<hr>
rubric={points:4}

Follow the [homework submission instructions](https://github.com/UBC-CS/cpsc330-2023s/blob/main/docs/homework_instructions.md). 

**You may work with a partner on this homework and submit your assignment as a group.** Below are some instructions on working as a group.  
- The maximum group size is 2. 
- Use group work as an opportunity to collaborate and learn new things from each other. 
- Be respectful to each other and make sure you understand all the concepts in the assignment well. 
- It's your responsibility to make sure that the assignment is submitted by one of the group members before the deadline. 
- If the .ipynb file is too big and doesn't render on Gradescope, also upload a pdf or html in addition to the .ipynb.
- You can find the instructions on how to do group submission on Gradescope [here](https://help.gradescope.com/article/m5qz2xsnjy-student-add-group-members).
- Note that we are using the autograder for some of the coding questions so that you can get instant feedback. When you submit your homework on Gradescope, it will run your code on an AWS server. It's possible that all your tests pass locally but some test cases fail when you submit your work on Gradescope. Wait for your submission to run there and examine the autograder results.  If it fails, you are either submitting a wrong file or doing something unexpected somewhere in the code. If you cannot figure it out, make use of office hours and tutorials. It's important that you start early on the assignments so that we help you with such issues. Also, remember that passing the tests is not the only goal. Make sure you are understanding what exactly you are doing in each question and why you are doing it. 

_Note: The assignments will get gradually more open-ended as we progress through the course. In many cases, there won't be a single correct solution. Sometimes you will have to make your own choices and your own decisions (for example, on what parameter values to use when they are not explicitly provided in the instructions). Use your own judgment in such cases and justify your choices, if necessary._

<br><br><br><br>

## Exercise 1: implementing `DummyClassifier`
<hr>
rubric={autograde:10}

In this course (unlike CPSC 340) you will generally **not** be asked to implement machine learning algorihtms (like logistic regression) from scratch. However, this exercise is an exception: you will implement the simplest possible classifier, `DummyClassifier`.
 
As a reminder, `DummyClassifier` is meant as a baseline and is generally a simple "model" you could "fit" to a dataset. All it does is predict the most popular class in the training set. So if there are more 0s than 1s it predicts 0 every time, and if there are more 1s than 0s it predicts 1 every time. For `predict_proba` it looks at the frequencies in the training set, so if you have 30% 0's 70% 1's it predicts `[0.3 0.7]` every time. Thus, `fit` only looks at `y` (not `X`).

Below you will find starter code for a class called `MyDummyClassifier`, which has methods `fit()`, `predict()`, `predict_proba()` and `score()`. Your task is to fill in those four functions. To get your started, I have given you a `return` statement in each case that returns the correct data type: `fit` returns nothing, `predict` returns an array whose size is the number of examples, `predict_proba` returns an array whose size is the number of examples x 2, and `score` returns a number.

The next code block has some tests you can use to assess whether your code is working. 

I suggest starting with `fit` and `predict`, and making sure those are working before moving on to `predict_proba`. For `predict_proba`, you should return the frequency of each class in the training data, which is the behaviour of `DummyClassifier(strategy='prior')`. Your `score` function should call your `predict` function. Again, you can compare with `DummyClassifier` using the code below.

To simplify this question, you can assume **binary classification**, and furthermore that these classes are **encoded as 0 and 1**. In other words, you can assume that `y` contains only 0s and 1s. The real `DummyClassifier` works when you have more than two classes, and also works if the target values are encoded differently, for example as "cat", "dog", "mouse", etc.

<div class="alert alert-warning">

Solution_1
    
</div>

_Points:_ 10

```python
class MyDummyClassifier:
    """
    A baseline classifier that predicts the most common class.
    The predicted probabilities come from the relative frequencies
    of the classes in the training data.

    This implementation only works when y only contains 0s and 1s.
    """

    most_freq = None
    zeros = 0
    ones = 0
    
    def fit(self, X, y):
        zeros = 0
        ones = 0
        for i in range(y.size):
            if y[i] == 0:
                zeros += 1
            elif y[i] == 1:
                ones += 1
        if zeros >= ones:
            most_freq = 0
        else:
            most_freq = 1
        return None

    def predict(self, X):
        if ones > zeros:
            return np.ones(X.shape[0])
        else:
            return np.zeros(X.shape[0])

    def predict_proba(self, X):
        total = zeros + ones
        return np.zeros([zeros / total, ones / total])  # Replace with your code

    def score(self, X, y):
        correct = 0
        return 0.0  # Replace with your code
    
```

```python
from sklearn.dummy import DummyClassifier
n_train = 21
d = 4
X_train_dummy = np.random.randn(n_train, d) 
y_train_dummy = np.random.randint(2, size=n_train)
my_dc = MyDummyClassifier()
my_dc.fit(X_train_dummy, y_train_dummy)
my_dc.predict(X_train_dummy)
my_dc.predict_proba(X_train_dummy)
my_dc.score(X_train_dummy, y_train_dummy)
```

Output:
```
0.0
```

```python
sk_dc = DummyClassifier(strategy="prior") # sklearn
sk_dc.fit(X_train_dummy, y_train_dummy)
sk_dc.predict(X_train_dummy)
sk_dc.predict_proba(X_train_dummy)
sk_dc.score(X_train_dummy, y_train_dummy)
```

Output:
```
0.6190476190476191
```

```python
grader.check("q1")
```

Output:
```
q1 results:
    q1 - 1 result:
        ❌ Test case failed
        Trying:
            def test_it(MyDummyClassifier):
                import matplotlib.pyplot as plt
                import numpy as np
                import pandas as pd
                from sklearn.dummy import DummyClassifier
                n_train = 101
                n_valid = 21
                d = 5
                for _ in range(100):
                    X_train_dummy = np.random.randn(n_train, d) 
                    X_valid_dummy = np.random.randn(n_valid, d)
                    y_train_dummy = np.random.randint(2, size=n_train)
                    y_valid_dummy = np.random.randint(2, size=n_valid)
                    my_dc = MyDummyClassifier()
                    sk_dc = DummyClassifier(strategy="prior")
                    _ = my_dc.fit(X_train_dummy, y_train_dummy)
                    _ = sk_dc.fit(X_train_dummy, y_train_dummy)
                    assert np.array_equal(my_dc.predict(X_train_dummy), sk_dc.
```

<br><br><br><br>

## Exercise 2: Trump Tweets
<hr>

For the rest of this assignment we'll be working with a [dataset of Donald Trump's tweets](https://www.kaggle.com/austinreese/trump-tweets) as of June 2020. You should start by downloading the dataset. Unzip it and move the file `realdonaldtrump.csv` under the data directory in this folder. As usual, please do not submit the dataset when you submit the assignment. 

```python
tweets_df = pd.read_csv("data/realdonaldtrump.csv", index_col=0)
tweets_df.head()
```

Output:
```
File ~/miniconda3/envs/cpsc330/lib/python3.10/site-packages/pandas/io/parsers/readers.py:1735, in TextFileReader._make_engine(self, f, engine)
   1733     if "b" not in mode:
   1734         mode += "b"
-> 1735 self.handles = get_handle(
   1736     f,
   1737     mode,
   1738     encoding=self.options.get("encoding", None),
   1739     compression=self.options.get("compression", None),
   1740     memory_map=self.options.get("memory_map", False),
   1741     is_text=is_text,
   1742     errors=self.options.get("encoding_errors", "strict"),
   1743     storage_options=self.options.get("storage_options", None),
   1744 )
   1745 assert self.handles is not None
   1746 f = self.handles.handle

File ~/miniconda3/envs/cpsc330/lib/python3.10/site-packages/pandas/io/common.py:856, in get_handle(path_or_buf, mode, encoding, compression, memory_map, is_text, errors, storage_options)
    851 elif isinstance(handle, str):
    852     # Check whether the filename is to be opened in binary mode.

```

```python
tweets_df.shape
```

We will be trying to predict whether a tweet will go "viral", defined as having more than 10,000 retweets:

```python
y = tweets_df["retweets"] > 10_000
```

To make predictions, we'll be using only the content (text) of the tweet. 

```python
X = tweets_df["content"]
```

For the purpose of this assignment, you can ignore all the other columns in the original dataset.

<br><br>

<br><br>

<!-- BEGIN QUESTION -->

### 2.1 Ordering the steps
rubric={points:6}

Let's start by building a model using `CountVectorizer` and `LogisticRegression`. The code required to do this has been provided below, but in the wrong order. 

**Your tasks:**
1. Rearrange the lines of code to correctly fit the model and compute the cross-validation score. 
2. Add a short comment above each step to describe what the code is doing.

<div class="alert alert-warning">

Solution_2.1
    
</div>

_Points:_ 6

```python
# YOUR COMMENT HERE
countvec = CountVectorizer(stop_words="english")

# YOUR COMMENT HERE
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.6, random_state=123)

# YOUR COMMENT HERE
cross_val_results = pd.DataFrame(
    cross_validate(pipe, X_train, y_train, return_train_score=True)
)

# YOUR COMMENT HERE
pipe = make_pipeline(countvec, lr)

# YOUR COMMENT HERE
cross_val_results.mean()

# YOUR COMMENT HERE
lr = LogisticRegression(max_iter=1000, random_state=123)
```

```python
...
```

<!-- END QUESTION -->

<br><br>

### 2.2 Baseline
rubric={autograde:2}

**Your tasks:**

1. Run 5-fold cross-validation with `DummyClassifier` on this dataset. Store the mean cross-validation score in a variable called `dummy_cv_score` given below.

> You should be able to use `MyDummyClassifier` above. But sklearn will likely complain if you try to pass it to `cross_val_score` or `cross_validate`. 

<div class="alert alert-warning">

Solution_2.2
    
</div>

_Points:_ 2

```python
dummy_cv_score = None

...
```

```python
grader.check("q2.2")
```

<br><br>

<!-- BEGIN QUESTION -->

### 2.3 Discussion
rubric={points:4}

**Your tasks:**

1. Compare the `DummyClassifier` score to what you got from logistic regression above. Does logistic regression seem to be doing anything useful?
2. Above we averaged the scores from the 5 folds of cross-validation of logistic regression. Print out the 5 individual scores. Reminder: `sklearn` calls them `"test_score"` but they are really (cross-)validation scores. 
3. Are the 5 scores close to each other or spread far apart? (This is a bit subjective, answer to the best of your ability.)
4. How does the size of this dataset (number of rows) compare to [the cities dataset](https://github.com/UBC-CS/cpsc330-2023s/blob/main/lectures/data/canada_usa_cities.csv) we have been using in class? How does this relate to the different sub-scores from the 5 folds?

<div class="alert alert-warning">

Solution_2.3
    
</div>

_Points:_ 4

_Type your answer here, replacing this text._

```python
...
```

```python
...
```

```python
...
```

<!-- END QUESTION -->

<br><br>

### 2.4 Probability scores
rubric={autograde:4}

Here we train a logistic regression classifier on the entire training set: 

> Note: this is relying on the `pipe` variable from 2.1 - you'll need to redefine it if you overwrote that variable in between.

```python
pipe.fit(X_train, y_train);
```

**Your tasks:**

1. Using this model, find the tweet in the **test set** with the highest predicted probability of being viral. Store the tweet and the associated probability in the variables `tweet` and `prob`, respectively. 

> Reminder: you are free to reuse/adapt code from lecture. Please add in a small attribution, e.g. "From Lecture 7".

<div class="alert alert-warning">

Solution_2.4
    
</div>

_Points:_ 4

```python
tweet = None
prob = None

...
```

```python
grader.check("q2.4")
```

<br><br>

### 2.5 Coefficients
rubric={autograde:4}

We can extract the `CountVectorizer` and `LogisticRegression` objects from the `Pipeline` object as follows:

```python
vec_from_pipe = pipe.named_steps["countvectorizer"]
lr_from_pipe = pipe.named_steps["logisticregression"]
```

**Your tasks:**

Using these extracted components above, get the five words with the highest coefficients and 5 words with smallest coefficients. Store them as lists in `top_5_words` and `bottom_5_words` variables, respectively. 

<div class="alert alert-warning">

Solution_2.5
    
</div>

_Points:_ 4

```python
top_5_words = None
bottom_5_words = None

...
```

```python
grader.check("q2.5")
```

<br><br>

### 2.6 Running a cross-validation fold without sklearn tools 
rubric={autograde:8}

Sklearn provides a lot of useful tools like `make_pipeline` and `cross_validate`, which are awesome. But with these fancy tools it's also easy to lose track of what is actually happening under the hood. 

**Your tasks:**

1. Compute logistic regression's validation score on the first fold, that is, train on 80% and validate on 20% of the training data (`X_train`) without using sklearn `Pipeline` or `cross_validate` or `cross_val_score`. Store the score of the fold in a variable called `fold_score`. Recall that `cross_validation` in `sklearn` does not shuffle the data by default.    

You should start with the following `CountVectorizer` and `LogisticRegression` objects, as well as `X_train` and `y_train` (which you should further split with `train_test_split` and `shuffle=False`):

```python
countvec = CountVectorizer(stop_words="english")
lr = LogisticRegression(max_iter=1000, random_state=123)
```

> Meta-comment: you might be wondering why we're going into "implementation" here if this course is about _applied_ ML. In CPSC 340, we would go all the way down into `LogisticRegression` and understand how `fit` works, line by line. Here we're not going into that at all, but I still think this type of question (and Exercise 1) is a useful middle ground. I do want you to know what is going on in `Pipeline` and in `cross_validate` even if we don't cover the details of `fit`. To get into logistic regression's `fit` requires a bunch of math; here, we're keeping it more conceptual and avoiding all those prerequisites.

<div class="alert alert-warning">

Solution_2.6
    
</div>

_Points:_ 8

```python
fold_score = None

...
```

```python
grader.check("q2.6")
```

<br><br><br><br>

## Exercise 3: Hyperparameter optimization
<hr>

<!-- BEGIN QUESTION -->

### 3.1 Optimizing `max_features` of `CountVectorizer`
rubric={points:2}

The following code varies the `max_features` hyperparameter of `CountVectorizer` and makes a plot (with the x-axis on a log scale) that shows train/cross-validation scores vs. `max_features`. It also prints the results. 

**Your tasks:**
- Based on the plot/output, what value of `max_features` seems best? Briefly explain.

> The code may take a minute or two to run. You can uncomment the `print` statement if you want to see it show the progress.

```python
train_scores = []
cv_scores = []

max_features = [10, 100, 1000, 10_000, 100_000]

for mf in max_features:
    #     print(mf)    
    pipe = make_pipeline(CountVectorizer(stop_words="english", max_features=mf), LogisticRegression(max_iter=1000))
    cv_results = cross_validate(pipe, X_train, y_train, return_train_score=True)
    train_scores.append(cv_results["train_score"].mean())
    cv_scores.append(cv_results["test_score"].mean())

plt.semilogx(max_features, train_scores, label="train")
plt.semilogx(max_features, cv_scores, label="valid")
plt.legend()
plt.xlabel("max_features")
plt.ylabel("accuracy");
```

```python
pd.DataFrame({"max_features": max_features, "train": train_scores, "cv": cv_scores})
```

<div class="alert alert-warning">

Solution_3.1
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

```python
...
```

<!-- END QUESTION -->

<br><br>

<!-- BEGIN QUESTION -->

### 3.2 Optimizing `C` of `LogisticRegression`
rubric={points:2}

The following code varies the `C` hyperparameter of `LogisticRegression` and makes a plot (with the x-axis on a log scale) that shows train/cross-validation scores vs. `C`. 

**Your tasks:**

- Based on the plot, what value of `C` seems best? Briefly explain. 

> The code may take a minute or two to run. You can uncomment the `print` statement if you want to see it show the progress.

```python
train_scores = []
cv_scores = []

C_vals = 10.0 ** np.arange(-1.5, 2, 0.5)

for C in C_vals:
    #     print(C)
    pipe = make_pipeline(CountVectorizer(stop_words="english"), LogisticRegression(max_iter=1000, C=C))    
    cv_results = cross_validate(pipe, X_train, y_train, return_train_score=True)

    train_scores.append(cv_results["train_score"].mean())
    cv_scores.append(cv_results["test_score"].mean())

plt.semilogx(C_vals, train_scores, label="train")
plt.semilogx(C_vals, cv_scores, label="valid")
plt.legend()
plt.xlabel("C")
plt.ylabel("accuracy");
```

```python
pd.DataFrame({"C": C_vals, "train": train_scores, "cv": cv_scores})
```

<div class="alert alert-warning">

Solution_3.2
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

### 3.3 Hyperparameter optimization 
rubric={autograde:10}

Start with the pipeline `pipe` below.

**Your tasks:**
- Create a `GridSearchCV` object named `grid_search` to jointly optimize `max_features` of `CountVectorizer` and `C` of `LogisticRegression` across all the combinations of values we tried above. 
- What are the best values of `max_features` and `C_vals` according to your grid search? Store them in variables `best_max_features` and `best_C`, respectively.  
- Store the best score returned by the grid search in a variable called `best_score`. 

> The code might be a bit slow here. Setting `n_jobs=-1` should speed it up if you have a multi-core processor.

```python
pipe = make_pipeline(CountVectorizer(stop_words="english"), LogisticRegression(max_iter=1000, random_state=123))
```

<div class="alert alert-warning">

Solution_3.3
    
</div>

_Points:_ 8

```python
grid_search = None 
best_max_features = None
best_C = None
best_score = None

...
```

```python
grader.check("q3.3")
```

<br><br>

<!-- BEGIN QUESTION -->

### 3.4 Discussion 
rubric={points:4}

- Do the best values of hyperparameters found by Grid Search agree with what you found in 3.1 and 3.2? 
- Generally speaking, _should_ these values agree with what you found in parts  3.1 and 3.2? Why or why not? 

<div class="alert alert-warning">

Solution_3.4
    
</div>

_Points:_ 4

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br>

### 3.5 Test score
rubric={autograde:2}

**Your tasks:**
- Evaluate your final model on the test set. Store the test accuracy in the variable called `test_score`.

<div class="alert alert-warning">

Solution_3.5
    
</div>

_Points:_ 2

```python
test_score = None

...
```

```python
grader.check("q3.5")
```

<br><br>

<!-- BEGIN QUESTION -->

### 3.6 Discussion
rubric={points:2}

- How does your test accuracy compare to your validation accuracy? 
- If they are different: do you think this is because you "overfitted on the validation set", or simply random luck?

<div class="alert alert-warning">

Solution_3.6
    
</div>

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

<!-- BEGIN QUESTION -->

## Exercise 4: Very short answer questions
<hr>
rubric={points:6}

Each question is worth 2 points. Max 2 sentences per answer.

1. What is the problem with calling `fit_transform` on your test data with `CountVectorizer`? 
2. If you could only access one of `predict` or `predict_proba`, which one would you choose? Briefly explain.
3. What are two advantages of `RandomizedSearchCV` over `GridSearchCV`?

<div class="alert alert-warning">

Solution_4
    
</div>

_Points:_ 6

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br><br><br><br>

## Submission instructions 

**PLEASE READ:** When you are ready to submit your assignment do the following:

1. Run all cells in your notebook to make sure there are no errors by doing `Kernel -> Restart Kernel and Clear All Outputs` and then `Run -> Run All Cells`. 
2. Notebooks with cell execution numbers out of order or not starting from “1” will have marks deducted. Notebooks without the output displayed may not be graded at all (because we need to see the output in order to grade your work).
3. Upload the assignment using Gradescope's drag and drop tool. Check out this [Gradescope Student Guide](https://lthub.ubc.ca/guides/gradescope-student-guide/) if you need help with Gradescope submission. 

Congratulations on finishing the homework! 

![](img/eva-well-done.png)



```python

```

Tests for q1 (points: 10)
```
>>> 
>>> 
>>> def test_it(MyDummyClassifier):
...     import matplotlib.pyplot as plt
...     import numpy as np
...     import pandas as pd
...     from sklearn.dummy import DummyClassifier
...     n_train = 101
...     n_valid = 21
...     d = 5
...     for _ in range(100):
...         X_train_dummy = np.random.randn(n_train, d) 
...         X_valid_dummy = np.random.randn(n_valid, d)
...         y_train_dummy = np.random.randint(2, size=n_train)
...         y_valid_dummy = np.random.randint(2, size=n_valid)
...         my_dc = MyDummyClassifier()
...         sk_dc = DummyClassifier(strategy="prior")
...         _ = my_dc.fit(X_train_dummy, y_train_dummy)
...         _ = sk_dc.fit(X_train_dummy, y_train_dummy)
...         assert np.array_equal(my_dc.predict(X_train_dummy), sk_dc.predict(X_train_dummy)), "the prediction score is wrong on a random training set"
...         assert np.array_equal(my_dc.predict(X_valid_dummy), sk_dc.predict(X_valid_dummy)), "the prediction score is wrong a random validation set"
>>> 
>>> test_it(MyDummyClassifier)
```
success_message: Good job!
```
>>> 
>>> 
>>> def test_it(MyDummyClassifier):
...     import matplotlib.pyplot as plt
...     import numpy as np
...     import pandas as pd
...     from sklearn.dummy import DummyClassifier
...     n_train = 101
...     n_valid = 21
...     d = 5
...     for _ in range(100):
...         X_train_dummy = np.random.randn(n_train, d) 
...         X_valid_dummy = np.random.randn(n_valid, d)
...         y_train_dummy = np.random.randint(2, size=n_train)
...         y_valid_dummy = np.random.randint(2, size=n_valid)
...         my_dc = MyDummyClassifier()
...         sk_dc = DummyClassifier(strategy="prior")
...         _ = my_dc.fit(X_train_dummy, y_train_dummy)
...         _ = sk_dc.fit(X_train_dummy, y_train_dummy)
...         assert np.allclose(my_dc.predict_proba(X_train_dummy), sk_dc.predict_proba(X_train_dummy)), "predict_proba is wrong on a random training set"
...         assert np.allclose(my_dc.predict_proba(X_valid_dummy), sk_dc.predict_proba(X_valid_dummy)), "predict_proba is wrong on a random validation set"
>>> 
>>> test_it(MyDummyClassifier)
```
success_message: Good job!
```
>>> 
>>> 
>>> def test_it(MyDummyClassifier):
...     import matplotlib.pyplot as plt
...     import numpy as np
...     import pandas as pd
...     from sklearn.dummy import DummyClassifier
...     n_train = 101
...     n_valid = 21
...     d = 5
...     for _ in range(100):
...         X_train_dummy = np.random.randn(n_train, d) 
...         X_valid_dummy = np.random.randn(n_valid, d)
...         y_train_dummy = np.random.randint(2, size=n_train)
...         y_valid_dummy = np.random.randint(2, size=n_valid)
...         my_dc = MyDummyClassifier()
...         sk_dc = DummyClassifier(strategy="prior")
...         _ = my_dc.fit(X_train_dummy, y_train_dummy)
...         _ = sk_dc.fit(X_train_dummy, y_train_dummy)
...         assert np.isclose(my_dc.score(X_train_dummy, y_train_dummy), sk_dc.score(X_train_dummy, y_train_dummy)), "the score is wrong on a random training set"
...         assert np.isclose(my_dc.score(X_valid_dummy, y_valid_dummy), sk_dc.score(X_valid_dummy, y_valid_dummy)), "the score is wrong on a random validation set"
>>> 
>>> test_it(MyDummyClassifier)
```
success_message: Good job!


Tests for q2.2 (points: 2)
```
>>> assert not dummy_cv_score is None, 'Are you using the provided variable?'
>>> assert sha1(str(np.round(dummy_cv_score, 3)).encode('utf-8')).hexdigest() == '838befe9ffa0a0d530805ba36c2e4890d4148ba1', "Your mean CV score doesn't look correct."
```


Tests for q2.4 (points: 4)
```
>>> assert not tweet is None, "Are you using the correct variable to store the tweet?"
>>> assert not prob is None, "Are you using the correct variable to store the probability?"
>>> assert sha1(str(np.round(prob, 4)).encode('utf-8')).hexdigest() == 'e8dc057d3346e56aed7cf252185dbe1fa6454411', "Incorrect probability."
>>> assert sha1(str(tweet).encode('utf8')).hexdigest() == 'fe39d5cbae2b335b4fde5486a8df189e41392043', "Incorrect tweet text."
```


Tests for q2.5 (points: 4)
```
>>> assert not top_5_words is None, "Are you using the correct variable?"
>>> assert not bottom_5_words is None, "Are you using the correct variable?"
>>> assert len(top_5_words) == 5, "Are you getting the top 5 words?"
>>> assert len(bottom_5_words) == 5, "Are you getting the bottom 5 words?"
>>> assert sha1("".join(sorted(top_5_words)).encode('utf-8')).hexdigest() == '439a61ad61c34c9069de01e2daa3a5bee63ee273', 'incorrect top 5 words'
>>> assert sha1("".join(sorted(bottom_5_words)).encode('utf-8')).hexdigest() == '5e2af1c788307cca183334510cf6bbee0208cab2', 'incorrect bottom 5 words'
```


Tests for q2.6 (points: 8)
```
>>> assert not fold_score is None, "Are you using the correct variable?"
>>> assert sha1(str(np.round(fold_score, 4)).encode('utf8')).hexdigest() == 'd267e2ea9b801d9f7ac454714a7fe6320075c255', "The fold_score doesn't look right."
```


Tests for q3.3 (points: [5, 1, 1, 1])
```
>>> assert not best_max_features is None, "Are you using the correct variable?"
>>> assert not best_C is None, "Are you using the correct variable?"
>>> assert type(grid_search.get_params()['estimator']) == Pipeline, "Are you passing a pipeline to the GridSearch?"
>>> assert sha1(str(grid_search.get_params()['param_grid']['countvectorizer__max_features']).encode('utf-8')).hexdigest() == 'a685d9f6a6425aa50a8a264bc5c467e7e5dbe781', "Are you using the max_features values from 3.1?"
>>> assert sha1(str(np.around(grid_search.get_params()['param_grid']['logisticregression__C'], 4)).encode('utf-8')).hexdigest() == 'd235bbfc65b1d312a1aaa7a27eae4c110d5b1d9b', "Are you using the C values from 3.2?"
```
```
>>> assert sha1(str(best_max_features).encode('utf-8')).hexdigest() == '8a12a315082a345f1a9d3ad14b214cd36d310cf8', "Best max feature doesn't seem correct."
```
```
>>> assert sha1(str(np.round(best_C, 4)).encode('utf-8')).hexdigest() == 'b1dfb61e10561718d8bfe529db2fe5827d62f164', "Best C doesn't seem correct."
```
```
>>> assert sha1(str(np.round(best_score, 4)).encode('utf-8')).hexdigest() == '567b99788b50aef4a1ebe4338c49910c4b5363fc', "Best score doesn't seem correct."
```


Tests for q3.5 (points: 2)
```
>>> assert not test_score is None, "Are you storing the score in the provided variable?"
>>> assert sha1(str(np.round(test_score, 4)).encode('utf-8')).hexdigest() == 'dad0d8c6709453d08f527227df265bb497d5067e', "The test score doesn't look correct."
```
